# Homework 2 (40 points)

name: Jiadong Zhu

email: jiadongz@umich.edu

This homework assignment took me 5 hours in total to complete. (Please help us to gauge the difficulty of the assignment.)

## Collaboration Disclosure

In the cell below, please list *everyone* with whom you discussed any of the homework problems, excluding only the GSIs and the course instructor. 

If you did not discuss the homework with anyone else, write __"I did not discuss this homework with anyone."__

Even if you discuss questions with other, the code you submit must be only yours. All work is checked with the [MOSS plagiarism detector](https://theory.stanford.edu/~aiken/moss/).

I did not discuss this homework with anyone.

## Submission Instructions
Your homework solutions should be written entirely in this Jupyter notebook file. Once it contains your solutions, you should submit this notebook through Canvas. Before submitting, please make sure make sure __Cells->Run__ executes without errors; errors in your code cause problems for our auto-grader. Thus, even though we may ask you to check for errors in your functions, you should not include in your submission any examples of your functions actually raising those errors.

## Nbgrader

We will be using `nbgrader` to grade your jupyter notebook. Additional tests will be used at grading time, you are encouraged to read the problem carefully and verify your code covers all appropriate cases.

**Be careful:** If a jupyter notebook cell takes longer than `60s` to run, the autograder will not grade it and you will receive zero credit for that question.

## Error checking

You do not need to do error checking (raising errors, etc.) in your code unless we explicitly ask you to so in a problem. <b>In this homework, if there is an element is an incorrect type within a correct type, you can raise either a ValueError or a TypeError</b>. For example, if a question asks for a tuple of ints and the input tuple is `(2,2,'s')`, either error is acceptable.Note: The tests shown are not comprehensive, so you are encouraged to read the problem carefully and verify your code covers all possible cases.
Homework tips
Start early! If you run into trouble installing things or importing packages, it’s best to find those problems well in advance, not the night before your assignment is due when we cannot help you!

Make sure you back up your work! At a minimum, do your work in a Dropbox folder. Better yet, use git, which is well worth your time and effort to learn.

Be careful to follow directions! Remember that Python is case sensitive. If you are ask you to define a function called my_function and you define a function called My_Function, you will not receive full credit. You may want to copy-paste the function names below to make sure that the functions in your notebook match.



### Problem 1: Severance textbook exercises (11 points)

#### (1 point) Chapter 6, Exercise 5: Use find and string slicing to extract the portion of the string after the colon character and then use the float function to convert the extracted string into a floating point number.

<!-- grade_id: 6_5, points: 0, grade: false, solution: true -->
```python
# Finish the question based on the string below:
string_example = 'X-DSPAM-Confidence:  0.8475'


# You final result should be defined as number, i.e.
# number = ?
# YOUR CODE HERE
colon_pos = string_example.find(':')
number_str = string_example[colon_pos + 1:].strip()
number = float(number_str)
print(number)
```

Output:
```
0.8475
```

<!-- grade_id: cell-7bdf8410ad9ecabc, points: 1, grade: true, solution: false -->
```python

```

#### (1 point) Chapter 7, Exercises 1: Write a program to read through a file and print the contents of the file (line by line) all in upper case. Executing the program for the first 5 lines of *mbox-short.txt* will look as follows:

FROM STEPHEN.MARQUARD@UCT.AC.ZA SAT JAN  5 09:14:16 2008

RETURN-PATH: \<POSTMASTER@COLLAB.SAKAIPROJECT.ORG\>

RECEIVED: FROM MURDER (MAIL.UMICH.EDU \[141.211.14.90\])

> BY FRANKENSTEIN.MAIL.UMICH.EDU (CYRUS V2.3.8) WITH LMTPA;
     
> SAT, 05 JAN 2008 09:14:16 -0500

<!-- grade_id: cell-db52ee6479e35ee3, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
file_name = 'mbox-short.txt'
f = open(file_name, "r")

line_count = 0
for line in f:
    print(line.upper().strip())  
    line_count += 1
    if line_count >= 5:
        break

f.close()
```

Output:
```
FROM STEPHEN.MARQUARD@UCT.AC.ZA SAT JAN  5 09:14:16 2008
RETURN-PATH: <POSTMASTER@COLLAB.SAKAIPROJECT.ORG>
RECEIVED: FROM MURDER (MAIL.UMICH.EDU [141.211.14.90])
BY FRANKENSTEIN.MAIL.UMICH.EDU (CYRUS V2.3.8) WITH LMTPA;
SAT, 05 JAN 2008 09:14:16 -0500
```

<!-- grade_id: cell-a1bc8449f489f8dc, points: 1, grade: true, solution: false -->
```python
try:
    open(file_name)
except FileNotFoundError:
    raise AssertionError("The file cannot be opened.")
```

#### (2 points) Chapter 7, Exercises 2: Write a program that uses (or takes in if you decide to write a function) a file name, and then read through the file and look for lines of the form:

X-DSPAM-Confidence: 0.8475

#### When you encounter a line that starts with “X-DSPAM-Confidence:” pull apart the line to extract the floating-point number on the line. Count these lines and then compute the total of the spam confidence values from these lines. When you reach the end of the file, print out the average spam confidence.

<!-- grade_id: X-DSPAM, points: 0, grade: false, solution: true -->
```python
file2 = 'mbox-short.txt'
file = open(file2, 'r')

total_confidence = 0
count = 0

for line in file:
    if line.startswith("X-DSPAM-Confidence:"):
        confidence_value = float(line.strip().split(":")[1])
        total_confidence += confidence_value
        count += 1
        
average_confidence = total_confidence / count
print(f"Average spam confidence: {average_confidence}")

file.close()
```

Output:
```
Average spam confidence: 0.7507185185185187
```

<!-- grade_id: cell-19ec1ef61fa7331f, points: 2, grade: true, solution: false -->
```python
try:
    open(file2)
except FileNotFoundError:
    raise AssertionError("The file cannot be opened.")
```

#### Chapter 8, Exercises 2: 
fhand = open('mbox-short.txt')


count = 0


for line in fhand:
    words = line.split()
    # print 'Debug:', words
    if len(words) == 0 : continue
    if words[0] != 'From' : continue
    print words[2]

#### (2 points) Figure out which line of the above program is still not properly guarded. See if you can construct a text file which causes the program to fail and then modify the program so that the line is properly guarded and test it to make sure it handles your new text file.  You do not need to include your text file in your submission, and you should comment out any lines of code which refer to this file.

user_data = "mbox-short.txt"


fhand = open(user_data, 'r')


count = 0

<b>The start of the for loop would be line 1 for the next question.</b><br>
```for line in fhand:
    words = line.split()
    # print 'Debug:', words
    if len(words) == 0 : continue
    if words[0] != 'From' : continue
    print words[2]```

__If the first line of the for loop line number 1, which line of the above program is still not properly guarded? Leave your answer as an integer.

<!-- grade_id: cell-4aa7d1f6e2246480, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
ans = 'The 5th line of the program is not properly guarded.'
```

<!-- grade_id: cell-05cd887839065f23, points: 1, grade: true, solution: false -->
```python
# Make sure that you assigned your int answer to the variable unguarded_line.
```

<!-- grade_id: cell-7770687564fc26ec, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
file_name = "mbox-short.txt"
fhand = open(file_name, 'r')

count = 0

for line in fhand:
    words = line.split()
    # print 'Debug:', words
    if len(words) == 0: continue
    if words[0] != 'From': continue
    if len(words) < 3: continue
    print(words[2])

fhand.close()
```

Output:
```
Sat
Fri
Fri
Fri
Fri
Fri
Fri
Fri
Fri
Fri
Fri
Fri
Fri
Fri
Fri
Fri
Fri
Fri
Fri
Fri
Fri
Thu
Thu
Thu
Thu
Thu
Thu
```

<!-- grade_id: cell-00d1e2a3da84a79a, points: 1, grade: true, solution: false -->
```python
try:
    open(file_name)
except FileNotFoundError:
    raise AssertionError("The file cannot be opened.")
```

#### (1 point) Chapter 8, Exercises 3: Rewrite the guardian code in the above example without two $if$ statements. Instead, use a compound logical expression using the $or$ logical operator with a single if statement.

<!-- grade_id: cell-f9593bd8e0e0bb9b, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
file_name = "mbox-short.txt"
fhand = open(file_name, 'r')

count = 0
for line in fhand:
    words = line.split()
    # print 'Debug:', words
    if len(words) == 0 or words[0] != 'From' or len(words) < 3:
        continue
    print(words[2])

fhand.close()
```

Output:
```
Sat
Fri
Fri
Fri
Fri
Fri
Fri
Fri
Fri
Fri
Fri
Fri
Fri
Fri
Fri
Fri
Fri
Fri
Fri
Fri
Fri
Thu
Thu
Thu
Thu
Thu
Thu
```

<!-- grade_id: cell-2953b11558c103c8, points: 1, grade: true, solution: false -->
```python
try:
    open(file_name)
except FileNotFoundError:
    raise AssertionError("The file cannot be opened.")
```

#### (1 point) Chapter 8, Exercises 5: Write a program to read through the mail box data and when you find a line that starts with “From”, you will split the line into words using the split function. We are interested in who sent the message, which is the second word on the From line.
From stephen.marquard@uct.ac.za Sat Jan 5 09:14:16 2008

#### You will parse the From line and print out the second word for each From line, then you will also count the number of From (not From:) lines and print out a count at the end.
This is a good sample output with a few lines removed:


python fromcount.py


Enter a file name: mbox-short.txt


stephen.marquard@uct.ac.za


louis@media.berkeley.edu


zqian@umich.edu

<!-- grade_id: cell-62e3962a1fb2bf0a, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
file_name = 'mbox-short.txt'
f = open(file_name, "r")

count = 0

for line in f:
    if line.startswith('From') and not line.startswith('From:'):
        space_1 = line.find(' ')
        space_2 = line.find(' ', space_1 + 1)
        print(line[space_1 + 1:space_2])
        count += 1
        
average_confidence = total_confidence / count
print(f"Count: {count}")

f.close()
```

Output:
```
stephen.marquard@uct.ac.za
louis@media.berkeley.edu
zqian@umich.edu
rjlowe@iupui.edu
zqian@umich.edu
rjlowe@iupui.edu
cwen@iupui.edu
cwen@iupui.edu
gsilver@umich.edu
gsilver@umich.edu
zqian@umich.edu
gsilver@umich.edu
wagnermr@iupui.edu
zqian@umich.edu
antranig@caret.cam.ac.uk
gopal.ramasammycook@gmail.com
david.horwitz@uct.ac.za
david.horwitz@uct.ac.za
david.horwitz@uct.ac.za
david.horwitz@uct.ac.za
stephen.marquard@uct.ac.za
louis@media.berkeley.edu
louis@media.berkeley.edu
ray@media.berkeley.edu
cwen@iupui.edu
cwen@iupui.edu
cwen@iupui.edu
Count: 27
```

<!-- grade_id: cell-972bcb771fa1da2f, points: 1, grade: true, solution: false -->
```python
try:
    open(file_name)
except FileNotFoundError:
    raise AssertionError("The file cannot be opened.")
```

#### (1 point) Chapter 9, Exercises 3: Write a program to read through a mail log, build a histogram using a dictionary to count how many messages have come from each email address, and print the dictionary.

<!-- grade_id: cell-0111fc1a5d93c960, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
file_name = 'mbox-short.txt'
f = open(file_name, "r")

email_counts = {}

for line in f:
    if line.startswith('From') and not line.startswith('From:'):
        space_1 = line.find(' ')
        space_2 = line.find(' ', space_1 + 1)
        email = line[space_1 + 1:space_2]
        
        email_counts[email] = email_counts.get(email, 0) + 1

for email, count in email_counts.items():
    print(f'{email}: {count}')
    
print('')
print(email_counts)

f.close()
```

Output:
```
stephen.marquard@uct.ac.za: 2
louis@media.berkeley.edu: 3
zqian@umich.edu: 4
rjlowe@iupui.edu: 2
cwen@iupui.edu: 5
gsilver@umich.edu: 3
wagnermr@iupui.edu: 1
antranig@caret.cam.ac.uk: 1
gopal.ramasammycook@gmail.com: 1
david.horwitz@uct.ac.za: 4
ray@media.berkeley.edu: 1

{'stephen.marquard@uct.ac.za': 2, 'louis@media.berkeley.edu': 3, 'zqian@umich.edu': 4, 'rjlowe@iupui.edu': 2, 'cwen@iupui.edu': 5, 'gsilver@umich.edu': 3, 'wagnermr@iupui.edu': 1, 'antranig@caret.cam.ac.uk': 1, 'gopal.ramasammycook@gmail.com': 1, 'david.horwitz@uct.ac.za': 4, 'ray@media.berkeley.edu': 1}
```

<!-- grade_id: cell-d251268af54f0baa, points: 1, grade: true, solution: false -->
```python
try:
    open(file_name)
except FileNotFoundError:
    raise AssertionError("The file cannot be opened.")
```

#### (1 point) Chapter 9, Exercises 4: Add code to the above program to figure out who has the most messages in the file. After all the data has been read and the dictionary has been created, look through the dictionary using a maximum loop (see Chapter 5: Maximum and minimum loops) to find who has the most messages and print how many messages the person has.

<!-- grade_id: cell-1567629def5bd621, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
file1 = 'mbox-short.txt'
f = open(file1, "r")

email_counts = {}

for line in f:
    if line.startswith('From') and not line.startswith('From:'):
        space_1 = line.find(' ')
        space_2 = line.find(' ', space_1 + 1)
        email = line[space_1 + 1:space_2]
        
        email_counts[email] = email_counts.get(email, 0) + 1

max_email = ''
max_count = 0

for email, count in email_counts.items():
    if count > max_count:
        max_email = email
        max_count = count

# Print the result
print(f"The email with the most messages is {max_email} with {max_count} messages.")

f.close()
```

Output:
```
The email with the most messages is cwen@iupui.edu with 5 messages.
```

<!-- grade_id: cell-ebfd8a271904d30d, points: 1, grade: true, solution: false -->
```python
try:
    open(file1)
except FileNotFoundError:
    raise AssertionError("The file cannot be opened.")
```

#### (1 point) Chapter 10, Exercises 1: Read and parse the “From” lines and pull out the addresses from the line. Count the number of messages from each person using a dictionary. After all the data has been read, print the person with the most commits by creating a list of (count, email) tuples from the dictionary. Then sort the list in reverse order and print out the person who has the most commits (you'll print the person with the most commits twice).

<!-- grade_id: cell-a8164b4b9647cf09, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-8573424abb14baa3, points: 1, grade: true, solution: false -->
```python
try:
    open(file2)
except FileNotFoundError:
    raise AssertionError("The file cannot be opened.")
```

### Problem 2: Fun with Strings (6 points)
1. <b>(2 points)</b>  A palindrome is a word or phrase that reads the same backwards and forwards. Write a function called `is_palindrome`, which takes a string as its only argument, and returns a Boolean. Your function should ignore spaces and capitalization in assessing whether or not a string is a palindrome, so that "tacocat" and "T A C O cat" are both considered palindromes.

<!-- grade_id: hw2-2_1, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
def is_palindrome(x):
    cleaned_str = ''.join([char.lower() for char in x if char.isalnum()])
    if cleaned_str == cleaned_str[::-1]:
        return True
    return False
```

<!-- grade_id: cell-6f018e99d2dda393, points: 0.5, grade: true, solution: false -->
```python
assert type(is_palindrome('level')) == bool
```

<!-- grade_id: palindrome, points: 1, grade: true, solution: false -->
```python
assert is_palindrome('level') == True
assert is_palindrome('WAS it a car or a cat I saw') == True
```

<!-- grade_id: space, points: 0.5, grade: true, solution: false -->
```python
assert is_palindrome('a') == True
```

2. <b>(2 points)</b>  Let us say that a word is “abecedarian” if its letters appear in alphabetical order(repeated  letters  are  okay).   So,  for  example,  “adder”  and  “beet”  are  abecedar-ian,  whereas  “dog”  and  “cat”  are  not.   Write  a  function `is_abecedarian`,  which takes a single argument in the form of a string and returns True if the argumentis abecedarian and False otherwise.  Here you may assume that the input consistsonly  of  alphabetic  characters  and  spaces.   Your  function  should  ignore  spaces,  so that the string "abcd efgh xyz" is considered abecedarian.  Your function should ignore capitalization, so that "aA","aa" and "aA" are all considered abecedarian.

<!-- grade_id: hw2-2_2, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
def is_abecedarian(x):
    cleaned_str = ''.join([char.lower() for char in x if char.isalnum()])
    i = 0
    while i < len(cleaned_str):
        return cleaned_str[i] < cleaned_str[i + 1]
        i += 1
```

<!-- grade_id: abc_blank, points: 0.5, grade: true, solution: false -->
```python
assert type(is_abecedarian("Ca T")) == bool
```

<!-- grade_id: abc-basic, points: 0.5, grade: true, solution: false -->
```python
assert is_abecedarian("Ca T") == False
```

<!-- grade_id: abc, points: 0.5, grade: true, solution: false -->
```python
assert is_abecedarian("adder") == True
```

<!-- grade_id: abc-caps, points: 0.5, grade: true, solution: false -->
```python
assert is_abecedarian("zYx") == False
```

3. <b>(2 points)</b> Write a function called `double_vowels` that takes a string as its only argument and returns that string with all the vowels duplicated.

<!-- grade_id: hw2-2_3, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
def double_vowels(x):
    vowels = 'aeiou'
    new_str = ''  # Initialize an empty string to store the result
    i = 0
    
    while i < len(x):
        if x[i] in vowels:
            new_str += x[i] * 2  # Double the vowel and add it to the new string
        else:
            new_str += x[i]  # Add the non-vowel character as it is
        i += 1
    
    return new_str
```

<!-- grade_id: spaces, points: 1, grade: true, solution: false -->
```python
assert double_vowels('this should work') == 'thiis shoouuld woork'
```

<!-- grade_id: double_vowels, points: 1, grade: true, solution: false -->
```python
assert double_vowels('vowel') == "vooweel"
```

### Problem 3: Fun with Lists (5 points)
1. <b>(2 points)</b> Write a function `list_reverse` that takes a list as an argument and returns that list, reversed. Your function should perform error checking to ensure that the input is a list. 

<!-- grade_id: hw2-3_1, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
def list_reverse(x):
    if not isinstance(x, list):
        raise TypeError("Input must be a list")
        
    return x[::-1]
```

<!-- grade_id: cell-1deca1a3bac64b18, points: 0.5, grade: true, solution: false -->
```python
assert list_reverse (['this should be the same list']) == ['this should be the same list']
```

<!-- grade_id: error, points: 1, grade: true, solution: false -->
```python
assert type(list_reverse([1,2,3,9])) == list
```

<!-- grade_id: list_reverse, points: 0.5, grade: true, solution: false -->
```python
assert list_reverse([1, 2, 3]) == [3, 2, 1]
```

2. <b>(1 point)</b> Write a function `is_sorted` that takes a sequence `seq` as its only argument and returns True if the sequence is sorted in ascending order and returns False otherwise.

<!-- grade_id: hw2-3_2, points: 0, grade: false, solution: true -->
```python
def is_sorted(seq):
    # YOUR CODE HERE
    i = 0
    while i < len(seq):
        return seq[i] < seq[i + 1]
        i += 1
```

<!-- grade_id: is_sorted, points: 1, grade: true, solution: false -->
```python
assert is_sorted('abbde') == True
assert is_sorted([1, 2, 3, 3, 5]) == True
```

3. <b>(2 points)</b> Write a function called `binary_search` that takes two arguments, a list of integers `t`(which is guaranteed to be sorted in ascending order) and an integer `elmt`, and returns True if `elmt` appears in list `t` and False otherwise. The binary search algorithm works as follows: 

    Search a sorted array by repeatedly dividing the search interval in half. Begin with an interval covering the whole array. If the value of the search key is less than the item in the middle of the interval, narrow the interval to the lower half. Otherwise narrow it to the upper half. Repeatedly check until the value is found or the interval is empty. (Source: https://www.geeksforgeeks.org/binary-search/)
    
    __Note:__ No credit will be given if binary search is not implemented. A simple for loop checking that the element is in the list will not qualify as a binary search.

<!-- grade_id: hw2-3_3, points: 0, grade: false, solution: true -->
```python
def binary_search(t, elmt):
    left, right = 0, len(t) - 1
    
    while left <= right:
        mid = (left + right) // 2
        if t[mid] == elmt:
            return True
        elif t[mid] > elmt:
            right = mid - 1
        else:
            left = mid + 1
    return False
```

<!-- grade_id: empty_list, points: 1, grade: true, solution: false -->
```python
assert type(binary_search([1,3], 0)) == bool
```

<!-- grade_id: binary_search, points: 1, grade: true, solution: false -->
```python
assert binary_search ([1, 3, 5, 7, 9], 9) == True
assert binary_search ([1, 3, 5, 7, 9], 1) == True
assert binary_search ([1, 3, 5, 7, 9], 7) == True
assert binary_search ([1, 3, 5, 7, 9], 0) == False
```

### Problem 4: More Fun with Strings (4 points)
1. <b>(2 points)</b> Write a function called `char_hist` that takes a string as its argument and returns a dictionary whose keys are characters and values are the number of times each character appeared in the input. The inputted string should be converted to have exclusively lowercase letters. 

<!-- grade_id: hw2-4_1, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
def char_hist(x):
    x = x.lower()
    char_count = {}

    for char in x:
        if char in char_count:
            char_count[char] += 1
        else:
            char_count[char] = 1
    
    return char_count
```

<!-- grade_id: empty_char, points: 1, grade: true, solution: false -->
```python
assert type(char_hist('This should return a dictionary.')) == dict
```

<!-- grade_id: char_hist, points: 1, grade: true, solution: false -->
```python
assert char_hist('gattaca') == {'g': 1, 'a': 3, 't': 2, 'c': 1}
```

2. <b>(2 points)</b> Write a function called `bigram_hist` that takes a string as its argument and returns a dictionary whose keys are 2-tuples of characters and values are the number of times that pair of characters appeared in the string. Your function should convert the string to have exclusive lowercase letters.

<!-- grade_id: hw2-4_2, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
def bigram_hist(x):
    x = x.lower()
    bigram_count = {}

    for i in range(len(x) - 1):
        bigram = (x[i], x[i+1])
        if bigram in bigram_count:
            bigram_count[bigram] += 1
        else:
            bigram_count[bigram] = 1
    
    return bigram_count
```

<!-- grade_id: bigram_hist, points: 1, grade: true, solution: false -->
```python
assert bigram_hist("mississippi") == {('i','p'):1,('i','s'):2,('m','i'):1,('p','i'):1,('p','p'):1,('s','i'):2,('s','s'):2}
```

<!-- grade_id: empty_hist, points: 1, grade: true, solution: false -->
```python
assert bigram_hist('') == {}
```

### Problem 5: Tuples as Vectors (6 points)
1. <b>(2 points)</b> Implement a function called `vec_scalar_mult`, which takes two arguments: a tuple of numbers (floats and/or integers) `t` and a number (float or integer) `s` and returns a tuple of the same length as `t`, with its entries equal to the entries of `t` multiplied by `s`. Your function should do appropriate error checking about the types of the arguments and the types within the vector. 

<!-- grade_id: hw2-5_1, points: 0, grade: false, solution: true -->
```python
def vec_scalar_mult(t, s):
    # YOUR CODE HERE
    if not isinstance(t, tuple):
        raise TypeError("The first argument must be a tuple.")

    if not isinstance(s, (int, float)):
        raise TypeError("The second argument must be a number (int or float).")

    for element in t:
        if not isinstance(element, (int, float)):
            raise TypeError("All elements of the tuple must be numbers (int or float).")

    return tuple(s * element for element in t)
```

<!-- grade_id: vec_scalar_mult, points: 1, grade: true, solution: false -->
```python
assert vec_scalar_mult ((1.0, 2), 4.0) == (4.0, 8.0)
```

<!-- grade_id: empty_mult, points: 1, grade: true, solution: false -->
```python
assert vec_scalar_mult ((1.0, 2), 4.0) == (4.0, 8.0)
assert vec_scalar_mult ((1.0, 2, 3), -2.0) == (-2.0, -4.0, -6.0)
```

2. <b>(2 points)</b> Implement a function called `vec_inner_product` which takes two vectors (i.e., tuples of foats and/or ints) as its inputs and outputs a float corresponding to the inner product of these two vectors. Your function should do appropriate error checking regarding the types and lengths of the arguments.

<!-- grade_id: hw2-5_2, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
def vec_inner_product(t1, t2):
    if not isinstance(t1, tuple) or not isinstance(t2, tuple):
        raise TypeError("Both arguments must be tuples.")

    if len(t1) != len(t2):
        raise ValueError("Both vectors must have the same length.")

    for element in t1 + t2:
        if not isinstance(element, (int, float)):
            raise TypeError("All elements of both vectors must be numbers (int or float).")

    inner_product = sum(x * y for x, y in zip(t1, t2))
    
    return float(inner_product)
```

<!-- grade_id: inner_product, points: 0.5, grade: true, solution: false -->
```python
assert vec_inner_product((1,1),(1,2)) == 3.0
assert vec_inner_product((1,1),(-1,2)) == 1.0
assert vec_inner_product((1.0, 2.0), (1, 2)) == 5.0
```

<!-- grade_id: inner-product-empty, points: 0.5, grade: true, solution: false -->
```python
assert vec_inner_product((), ()) == 0.0
```

<!-- grade_id: inner-pr, points: 0.5, grade: true, solution: false -->
```python
assert vec_inner_product((1,5,9,4.5), (7,-2,3,-4)) == 6.0
```

<!-- grade_id: error_handling, points: 0.5, grade: true, solution: false -->
```python
try:
    vec_inner_product((1.0, 2.0), (1.0,))

except ValueError:
    pass
except TypeError:
    pass

else:
    raise AssertionError("Dimensions don't line up.")
```

3. <b>(1 point)</b> Write a function `check_valid_mx` that takes a single argument and returns a Boolean, which is True if the given argument is a tuple that validly represents a matrix, and returns False otherwise. A valid matrix is a tuple of tuples that has exclusively numeric values (floats or ints). Additionally, each innermost tuple must have the same length. If the length of the first tuple is 5, then all other inner tuples must have length 5 in order for the structure to be considered a valid matrix. 

<!-- grade_id: hw2-5_3, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
def check_valid_mx(mx):
    for row in mx:
        if not all(isinstance(element, (int, float)) for element in row):
            return False

    row_length = len(mx[0]) if mx else 0
    if not all(len(row) == row_length for row in mx):
        return False
    
    return True
```

<!-- grade_id: check_valid_mx, points: 0.5, grade: true, solution: false -->
```python
assert check_valid_mx('s') == False
assert check_valid_mx(((1.0, 2.0), (2.0,))) == False
assert check_valid_mx(((1.0, 2.0), (2.0, 4.0))) == True
```

<!-- grade_id: cell-b66e8d0741a7c7c4, points: 0.5, grade: true, solution: false -->
```python
assert check_valid_mx(((1.0, 2.0), (2.0, 's'))) == False
```

4. <b>(1 point)</b> Write a function `mx_vec_mult` that takes a matrix (i.e., tuple of tuples) and a vector (i.e., a tuple) as its arguments, and returns a vector (i.e., a tuple of numbers) that is the result of multiplying the given vector by the given matrix (you are multiplying B(matrix)*a(vector)). Your function should check that all of the types of the parameters are correct and that they are all valid with matching dimensions for matrix-vector multiplication. __Hint:__ Think carefully about what type of error should be raised in each situation.

<!-- grade_id: hw2-5_4, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
def mx_vec_mult(mx, vec):
    num_cols = len(mx[0])
    if len(vec) != num_cols:
        raise ValueError("The number of columns in the matrix must equal the number of elements in the vector.")

    result = tuple(sum(row[i] * vec[i] for i in range(num_cols)) for row in mx)
    
    return result
```

<!-- grade_id: mx_vec_mult, points: 0.5, grade: true, solution: false -->
```python
assert mx_vec_mult(((1.0, 2.0), (2.0, 3.0)), (1.0, 2.0)) == (5.0, 8.0)
```

<!-- grade_id: cell-0bf79440f86eb75b, points: 0.5, grade: true, solution: false -->
```python
try:
    mx_vec_mult(((1.0, 2.0), (2.0, 's')), (1.0, 2.0))

except ValueError:
    pass
    
except TypeError:
    pass

else:
    raise AssertionError("The matrix is invalid.")
```

###  Problem 6: More Fun with Vectors (3 points)
1. <b>(1 point)</b> A sparse vector (or matrix) is one that is reporesented by only the storage of its non-zero entries, with (index, value) pairs. Write a function `is_valid_sparse_vector` that takes one argument, and returns True if and only if the input is a valid sparse vector, and returns False otherwise. Note: your function should not assume that the input is a dictionary. A valid sparse vector will be a dictionary that has the properties that 
    1. all its indices are positive integers, and 
    2. all its values are floats. 



<!-- grade_id: hw2-6_1, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
def is_valid_sparse_vector(sparse_vector):
    if not isinstance(sparse_vector, dict):
        return False
        
    for key, value in sparse_vector.items():
        if not isinstance(key, int) or key <= 0:
            return False
        if not isinstance(value, float):
            return False
            
    return True
```

<!-- grade_id: is_valid_sparse_vector, points: 0.5, grade: true, solution: false -->
```python
assert is_valid_sparse_vector({1:2.,6:1.,5:3.5,7:4.0}) == True
assert is_valid_sparse_vector({-2:1.2,3:2.5,5:4.1}) == False
assert is_valid_sparse_vector({1: 1.0, 3: 2.0, 5: 3.0, 7: 4.0}) == True
assert is_valid_sparse_vector({-1: 1.0, 3: 2.0, 5: 3.0}) == False
```

<!-- grade_id: is-valid-sparse-vector-empty, points: 0.5, grade: true, solution: false -->
```python
assert is_valid_sparse_vector({1}) == False
```

2. <b>(2 points)</b> Write a function `sparse_inner_product` that takes two sparse vectors as its inputs, and returns a float that is the value of the inner product of the vectors that the inputs represent. Note that by our definition, a sparse vector has no specified dimension, so there is no need to check that the dimensions of the arguments agree. Your function should do the appropriate error checking for ensuring that both inputs are valid sparse vectors.

<!-- grade_id: hw2-6_2, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
def sparse_inner_product(vec1, vec2):
    def is_valid_sparse_vector(sparse_vector):
        if not isinstance(sparse_vector, dict):
            return False
        for key, value in sparse_vector.items():
            if not isinstance(key, int) or key <= 0:
                return False
            if not isinstance(value, float):
                return False
        return True
        
    if not is_valid_sparse_vector(vec1) or not is_valid_sparse_vector(vec2):
        raise ValueError("Both inputs must be valid sparse vectors.")
        
    inner_product = 0.0
    for index, value in vec1.items():
        if index in vec2:
            inner_product += value * vec2[index]
    
    return inner_product
```

<!-- grade_id: cell-ab54723e53513a3c, points: 0.5, grade: true, solution: false -->
```python
assert sparse_inner_product({1: 1.0, 2: 2.0, 3: 3.0}, {1: 1.0, 3: 3.0, 5: 5.0}) == 10.0
```

<!-- grade_id: cell-cde1f885231cf4bd, points: 0.5, grade: true, solution: false -->
```python
assert sparse_inner_product({1: 1.0, 2: 2.0, 3: 3.0}, {}) == 0
```

<!-- grade_id: cell-e8d82c1c87291d19, points: 0.5, grade: true, solution: false -->
```python
assert sparse_inner_product({1: 1.0, 4: -3.5, 3: -3.5}, {1: 10.0, 3: 1.5, 5: 3.14}) == 4.75
```

<!-- grade_id: cell-d4d7e89ace52b2c4, points: 0.5, grade: true, solution: false -->
```python
assert sparse_inner_product({6: 1.0, 4: 2.0, 37: 3.0}, {2: 1.0, 1: 3.0, 55: 5.0}) == 0.0
```

### Problem 7: More Fun with Tuples (5 points)
1. <b>(2 points)</b> Write a function called `my_sum` that takes any number of numeric (ints and floats) arguments, and returns the sum of its arguments.

<!-- grade_id: hw2-7_1, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
def my_sum(*args):
    total = 0
    for n in args:
        total += n

    return total
```

<!-- grade_id: star-args, points: 1, grade: true, solution: false -->
```python
## Please make sure this cell is directly beneath the cell with your my_sum function.
```

<!-- grade_id: my_sum, points: 0.5, grade: true, solution: false -->
```python
assert my_sum(3.2,4.1,-3) == 4.3
assert my_sum(3,4) == 7
```

<!-- grade_id: cell-2c0a2ace639c9c61, points: 0.5, grade: true, solution: false -->
```python
assert my_sum(-0) == 0
```

2. <b>(1 point)</b> Write a function called `reverse_tuple` that takes a tuple as its only argument and returns a tuple that is the reverse of the input.

<!-- grade_id: hw2-7_2, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
def reverse_tuple(x):
    return x[::-1]
```

<!-- grade_id: cell-9aa079959f08d94e, points: 0.5, grade: true, solution: false -->
```python
assert type(reverse_tuple((4,5,6))) == tuple
```

<!-- grade_id: cell-66e2cb2b968800d1, points: 0.5, grade: true, solution: false -->
```python
assert reverse_tuple((3,4,1)) == (1,4,3)
assert reverse_tuple((4,1)) == (1,4)
```

3. <b>(2 points)</b> Write a function called `rotate_tuple` that takes two arguments: a tuple and an integer, in that order. Letting $n$ be the integer given in the input, your function should return a tuple of the same length as the input tuple, but with its entries “rotated” by n. If n is positive, this should mean to “push forward” all the entries of the input tuple by n entries, with entries that “go off the end” of the tuple being wrapped around to the beginning, so that the $i^{th}$ entry of the input tuple becomes the $(i + n)^{th}$ entry of the output, wrapping around to the beginning of the tuple if this index goes off the end. If n is negative, then this corresponds to rotating the entries in the other direction, with entries of the input tuple being “pushed backward”. Your function should perform error checking to ensure that the inputs are of appropriate types. If the user supplies a non-integer, print a message to alert the user that the input was not as expected, and try to recover by casting it to an integer. __Hint__: a try/catch statement will likely be useful here.

<!-- grade_id: hw2-7_3, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
def rotate_tuple(my_tuple, n):
    n = int(n)
    if n >= 0:
        return my_tuple[-n:] + my_tuple[:len(my_tuple) - n]
    else:
        return my_tuple[-n:] + my_tuple[:-n]
```

<!-- grade_id: floats_and_empty, points: 0.5, grade: true, solution: false -->
```python
assert type(rotate_tuple((1,2,3,4), 5)) == tuple
```

<!-- grade_id: neg_and_large_index, points: 1, grade: true, solution: false -->
```python
assert rotate_tuple((1,2,3,4,5,6,7),2) == (6,7,1,2,3,4,5)
assert rotate_tuple((1,2,3,4,5,6,7),-2) == (3,4,5,6,7,1,2)
```

<!-- grade_id: cell-8eefa6355b206775, points: 0.5, grade: true, solution: false -->
```python
assert rotate_tuple((1,2,3,4,5,6,7),-2.0) == (3,4,5,6,7,1,2)
```